# Functions

References:

[1] B Downey, A. (2012). *Think Python: How to Think Like a Computer Scientist-2e.*

[2] Gries, P., Campbell, J., & Montojo, J. (2017). *Practical programming: an introduction to computer science using Python 3.6.* Pragmatic Bookshelf.

[3] Matthes, E. (2023). *Python crash course: A hands-on, project-based introduction to programming.*

## 1 Functions in Python

One way to improve the modularity of your code is to adhere to the **DRY** principle, which means *Don't repeat yourself*. In other words, whenever you found yourself doing the same thing over and over again, do consider replacing it with abstraction or modules to avoid redundancy in your code.

### Defining functions in Python

Say for example we have a friend that lives in the other side of the world (United States), and you're always talking about the weather and how hot it is in your respective place. The US typically uses `Fahrenheit` as their unit of temperature measurement, thus, we typically need to convert fahrenheit to celcius and back a lot. It would be nice to be able to do convert these units easily given any number.

```python
def convert_to_fahrenheit(celsius):
    """Convert a given temperature in celsius to fahrenheit"""
    fahrenheit = 9/5*celsius + 32

    return fahrenheit

def convert_to_celsius(fahrenheit):
    """Convert a given temperature in fahrenheit to celsius"""
    celsius = (fahrenheit - 32) * 5/9

    return celsius
```

The above code block are examples of *function definitions* in Python. Here the *function body* which contains the code that will be executed upon function *call* is indented.

The first line of the function definition is called the *function header*. which contains the *function name* and its corresponding *arguments*. Arguments are expressions that appears between the parenthesis of a function that the function needs to execute the code block defined inside it. We input it upon *calling* the function.

The general form of a function definitoin is as follows:

```python
def <<function_name>>(<<arguments>>):
    <<function_body>>
```

Defining a function creates a function object, which has the type `function`

```python
type(convert_to_celsius)
```

You then call the function by using its name together with the required parameters enclosed in a parenthesis

```python
convert_to_celsius(98.6)
```

```python
convert_to_fahrenheit(29.0)
```

Most function definitions have a `return` statement, that when executed, ends the function and produces a *value*.

```python
return <<expression>>
```

When Python executes a `return` statement, it evaluates the expression then produces the result of that expression as the result of the function call.

### Passing arguments

When you pass arguments to a function, Python matches those arguments to the function's parameters depending on the way you specify them. You can pass arguments in a number of ways. You can use *positional arguments*, which need to be in the same order in which the parameters were written; *keyword arguments*, where each argument constists of a variable name and value.

```python
def describe_pet(animal_type, pet_name):
    """Display information about a pet."""
    print(f"I have a {animal_type}.")
    print(f"My {animal_type}'s name is {pet_name.title()}.")
```

```python
describe_pet('dog', 'Maki')
```

```python
describe_pet(animal_type='cat', pet_name='Midnight')
```

```python
describe_pet('Maki', 'dog')
```

```python
describe_pet(pet_name='Midnight', animal_type='cat')
```

Sometimes it makes sense to give an argument a default value, for example, in a scenario where you want to make the argument optional.

```python
def get_formatted_name(first_name, last_name, middle_name=''):
    """Return full name, neatly formatted written last name first."""
    if middle_name:
        full_name = f"{last_name}, {first_name} {middle_name}"
    else:
        full_name = f"{last_name}, {first_name}"

    return full_name.title()
```

```python
get_formatted_name("Leodegario", "Lorenzo")
```

```python
get_formatted_name("Leodegario", "Lorenzo", "Urgel")
```

### Recursive Functions

Recursive functions are functions that calls itself in the body. Knowing how to write recursive functions has a benefit of creating repetitions without the use of loops. However, be cautious when implementing recursive functions as you can easily make a recursive function that doesn't terminate itself, or one that uses an excessive amount of memory or processing power. Nonetheless, when used correctly, it allows for an elegant approach to programming.

An example of a recursive function is the factorial definition, which mathematically can be written as:

$$n! = n \times (n-1)!$$

```python
def factorial(n):
    """Return the factorial of integer value n"""
    if n == 0:
        return 1
    else:
        return n * factorial(n-1)
```

```python
factorial(4)
```

Another example is the definition of the Fibonacci numbers:

$$F_0 = 0, F_1 = 1$$
$$F_n = F_{n-1} + F_{n-2}$$

The first 6 Fibonacci numbers are:

$$F_0 = 0$$
$$F_1 = 1$$
$$F_2 = 1$$
$$F_3 = 2$$
$$F_4 = 3$$
$$F_5 = 5$$

```python
def fibonacci(n):
    """Return the fibonacci number F_n"""
    if n == 0:
        return 0
    elif n == 1:
        return 1
    else:
        return fibonacci(n-1) + fibonacci(n-2)
```

```python
fibonacci(5)
```

## 2 Designing New Functions: A Recipe

Everytime you write a function, do figure out the answers to the following questions:

1. What is an appropriate and informative name for the function?
2. What are the parameters, and what types of information do they refer to?
3. What calculations are you doing with that information?
4. What information does the function return?
5. Does it work like you expect it to?

Part of the outcome upon completing this questions should be a working function. But equally important is the *documentation* for the function. Use three double quotes `"""` to start and end this documentation. This notation is called the *docstring* or the *documentation string*.

Let's improve our temperature converter functions according to this framework.

*Note: the docstring format we specifically use is the [numpy](https://numpydoc.readthedocs.io/en/latest/format.html) docstring format. Which is the common notation for many data science libraries that you'll eventually be using*

```python
def convert_to_fahrenheit(celsius: float) -> float:
    """Convert a given temperature in celsius to fahrenheit

    Parameters
    ----------
    celsius : float
        Input temperature in celsius

    Returns
    -------
    fahrenheit : float
        Temperature of the input in fahrenheit

    Examples
    --------
    >>> convert_to_fahrenheit(37.0)
    98.6
    """
    fahrenheit = 9/5*celsius + 32

    return fahrenheit

def convert_to_celsius(fahrenheit: float) -> float:
    """Convert a given temperature in fahrenheit to celsius

    Parameters
    ----------
    fahrenheit : float
        Input temperature in fahrenheit

    Returns
    -------
    celsius : float
        Temperature of the input in celsius

    Examples
    --------
    >>> convert_to_celsius(98.6)
    37.0
    """
    celsius = (fahrenheit - 32) * 5/9

    return celsius
```

### Functions That Python Provides

To improve also the readability and conciseness of your code, consider using built-in functions to perform common operations. We've already encountered some of these functions in the previous session which includes `print`, `float`, `int`, `help`, and `input`.

Found [here](https://www.w3schools.com/python/python_ref_functions.asp) are other common built-in functions in Python: 

```python
abs(-9)
```

```python
round(1.337, 2)
```

```python
help(round)
```

```python
pow(25, 0.5)
```

## 3 A Modular Approach to Program Organization

> *Mathematicians don't prove every theorem from scratch. Instead, they build their proofs on truths their predecessors have already established. In the same way, it's rare for someone to write all of a program alone; it's much more common-and productive to make use of the millions of lines of code that other programmers have written before*

In Python, a *module* is a collection of variables and functions grouped together in a single file. These variables and functions usually help you accomplish a certain task. For example, the `math` module contains variables such as `pi` and mathematical functions such as `cos` (cosine) and `sqrt` (square root).

### Importing Modules

To gain access to the variables and functions defined in a module, we use the `import` statement.

#### `math` module

```python
import math
```

```python
type(math)
```

```python
help(math)
```

```python
math.sqrt(9)
```

## 4 Hands-on Exercises

### City names

Write a function called `city_country()` that takes in a name of a `city` and its `country`. The function should return a string formatted like this:

```
"Makati City, Philippines"
```

<!-- grade_id: cell-667d20c714a10725, points: 0, grade: false, solution: true -->
```python
def city_country(city, country):
    # YOUR CODE HERE
```

<!-- grade_id: cell-4359f1074a7acadd, points: 0, grade: true, solution: false -->
```python
assert isinstance(city_country('Makati City', 'Philippines'), str)
```

### Sum digits

Given a number `n`, create a function that computes the sum of the digits of this number.

<!-- grade_id: cell-69f87e694fdcb34e, points: 0, grade: false, solution: true -->
```python
def sum_digits(n):
    # YOUR CODE HERE
```

<!-- grade_id: cell-0f92ff5b7357b1a9, points: 0, grade: true, solution: false -->
```python
assert isinstance(sum_digits(1345), int)
```

### The Ackermann function

The [Ackermann function](https://en.wikipedia.org/wiki/Ackermann_function), $A(m, n)$ is defined as:

$$
A(m, n) = \begin{cases} n + 1 & \text{ if } m = 0 \\ A(m-1, 1) & \text{ if } m > 0 \text{ and } n = 0 \\ A(m-1, A(m, n-1)) & \text{ if } m > 0 \text{ and } n > 0 \end{cases}
$$

Write a function `ack` that evaluates the Ackermann function.

<!-- grade_id: cell-a1f035d85db41e73, points: 0, grade: false, solution: true -->
```python
def ack(m, n):
    # YOUR CODE HERE
```

<!-- grade_id: cell-f28bab47aa663cc0, points: 1, grade: true, solution: false -->
```python
assert isinstance(ack(3, 4), int)
```